# Assigment 3

### Instructions
- Do NOT change the filename.
- Do NOT add or copy cells before the "Sandpit". Only add your own cells after the "Sandpit".
- Enter your name below.

### Enter your name here
Name: CHANGE ME

# Task 2 (6 points)

**Important requirements**
- No loops are allowed for this task.
- The body of each function must be only **one** line of code.

Violating any of the requirements will score **ZERO** points for the corresponding function.

You are given the following dictionary of sets:

```python
prog_skills = {
    "Anna": {"C#", "Java"},
    "Bob": {"Python", "C++", "JavaScript"},
    "Carla": {"Java", "JavaScript", "Python"},
}
```

Using this dictionary, write a function that returns the answer to each of the questions below. The returned string should follow the format of the expected return value below.

### Question 1: What language(s) can a person code in?

For Carla, expected return value is:
``Carla can code in Java and JavaScript and Python.``

Hint: Google joining strings with a separator.

(Order of items in a set may vary, and that's just fine.)

<!-- grade_id: cell-c8c1768e7f527c52, points: 0, grade: false, solution: true -->
```python
def question1(dictionary, person):
    # YOUR CODE HERE
    raise NotImplementedError()
```

### Question 2: How many language(s) can a person code in?

For Bob, expected return value is:
``Bob codes in 3 languages.``

<!-- grade_id: cell-e67bd891d067ff8a, points: 0, grade: false, solution: true -->
```python
def question2(dictionary, person):
    # YOUR CODE HERE
    raise NotImplementedError()
```

### Question 3: What language(s) do both person 1 and person 2 code in?

For Anna and Carla, expected return value is:
``They both code in Java.``

For Bob and Carla, expected return value is:
``They both code in Java and JavaScript.``

(Order of items in a set may vary, and that's just fine.)

<!-- grade_id: cell-628b0844864613bc, points: 0, grade: false, solution: true -->
```python
def question3(dictionary, person1, person2):
    # YOUR CODE HERE
    raise NotImplementedError()
```

# Tests

The cell below contains unit tests for your solution.

Run the below cell (after running all the above cells with your code).

<!-- grade_id: cell-b8fd170e6b42cf68, points: 1, grade: true, solution: false -->
```python
assert question1(prog_skills, "Anna") == "Anna can code in C# and Java." or \
question1(prog_skills, "Anna") == "Anna can code in Java and C#."
```

<!-- grade_id: cell-c94254769ab304cf, points: 1, grade: true, solution: false -->
```python
# Hidden tests
```

<!-- grade_id: cell-7aae16280978cf72, points: 1, grade: true, solution: false -->
```python
assert question2(prog_skills, "Bob") == "Bob codes in 3 languages."
```

<!-- grade_id: cell-47a2056cb5457422, points: 1, grade: true, solution: false -->
```python
# Hidden tests
```

<!-- grade_id: cell-53305c7f6d30921e, points: 1, grade: true, solution: false -->
```python
assert question3(prog_skills, "Anna", "Carla") == "They both code in Java."
```

<!-- grade_id: cell-0912717fff9bfef6, points: 1, grade: true, solution: false -->
```python
# Hidden tests
```

# Sandpit begins here
### Add your own cells below