# Assigment 3

### Instructions
- Do NOT change the filename.
- Do NOT add or copy cells before the "Sandpit". Only add your own cells after the "Sandpit".
- Enter your name below.

### Enter your name here
Name: CHANGE ME

# Task 1 (2 points)

Write a function named ``multiply`` that multiplies an arbitrarily long list of number arguments (not passed as a list/collection but as single arguments) and returns the result.

It's OK to use loop(s) for this task.

<!-- grade_id: cell-810457514e42bdba, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# Tests

The cell(s) below contains unit test(s) for your solution.

Run the below cell (after running all the above cells with your code).

<!-- grade_id: cell-336ad2bf9b243033, points: 1, grade: true, solution: false -->
```python
assert multiply(1, 2, 4) == 8
```

<!-- grade_id: cell-0dae3e9c4f7a894d, points: 1, grade: true, solution: false -->
```python
# Hidden tests
```

# Sandpit begins here
### Add your own cells below