# Assigment 3

### Instructions
- Do NOT change the filename.
- Do NOT add or copy cells before the "Sandpit". Only add your own cells after the "Sandpit".
- Enter your name below.

### Enter your name here
Name: CHANGE ME

# **Optional** Task 4 (5 bonus points)

**Important requirements**
- No loops are allowed for this task.

Violating any of the requirements will score **ZERO** points.

Write a function that returns, in a list, all the primes from min to a given max (inclusive) number.

You can assume that min is at least 2, i.e. we will never call the function with min = 1.

Here are some sample outputs:

```
    primes(2, 5) --> [2, 3, 5]
    primes(10, 15) --> [11, 13]
```

<!-- grade_id: cell-9517a5994bf47076, points: 0, grade: false, solution: true -->
```python
def primes(min, max):
    # YOUR CODE HERE
    raise NotImplementedError()
```

# Tests

The cell below contains unit tests for your solution.

Run the below cell (after running all the above cells with your code).

<!-- grade_id: cell-bff660f3431edc64, points: 1, grade: true, solution: false -->
```python
assert primes(2, 5) == [2, 3, 5]
```

<!-- grade_id: cell-173f5eb75ab4a0c4, points: 1, grade: true, solution: false -->
```python
assert primes(10, 15) == [11, 13]
```

<!-- grade_id: cell-2cf2537cb3df28a8, points: 1, grade: true, solution: false -->
```python
# Hidden tests
```

<!-- grade_id: cell-a3c12b13f3df42ae, points: 1, grade: true, solution: false -->
```python
# Hidden tests
```

<!-- grade_id: cell-f01be3ae91e772dd, points: 1, grade: true, solution: false -->
```python
# Hidden tests
```

# Sandpit begins here
### Add your own cells below

```python

```