# Assigment 2

### Instructions
- Do NOT change the filename.
- Do NOT add or copy cells before the "Sandpit". Only add your own cells after the "Sandpit".
- Enter your name below.

### Enter your name here
Name: CHANGE ME

## Task 1: String manipulation
###### (2 points)
- Write code that creates a BHT email address, in this format:
  first 2 letters in first name in lower case, followed by first 2 letters in last name in lower case, followed by "1234"
- Example: "Butter Ball" gets the email address "buba1234@bht-berlin.de"
- (We haven't talked about functions yet, but we need this construct here for auto-grading.)

<!-- grade_id: cell-9e08127856761ea8, points: 0, grade: false, solution: true -->
```python
def create_email_address(first_name, last_name):
    # final line of code should look like this:
    # return a_string_with_email_address
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-70f1705262d021b6, points: 1, grade: true, solution: false -->
```python
assert create_email_address("Butter", "Ball") == "buba1234@bht-berlin.de"
```

Output:
```
Cell In[2], line 1
----> 1 assert create_email_address("Butter", "Ball") == "buba1234@bht-berlin.de"

Cell In[1], line 5, in create_email_address(first_name, last_name)
      1 def create_email_address(first_name, last_name):
      2     # final line of code should look like this:
      3     # return a_string_with_email_address
      4     # YOUR CODE HERE
----> 5     raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-6602798f99dcb2b5, points: 1, grade: true, solution: false -->
```python
assert create_email_address("Rainbow", "Unicorn") == "raun1234@bht-berlin.de"
```

## Task 2: List manipulation
###### (2 points)
Write code the performs the following:
- Put the following strings/words into a list (your choice of variable name), exactly in this order:

``['the', 'dog',  'quick', 'lazy', 'brown', 'a', 'fox', 'over', 'jumped']``

- Using up to __3__ carefully and cleverly chosen list operations, come up with a new list with (variable name __must__ be "new_list") the same words, but in this new order:

``['the', 'quick', 'brown', 'fox', 'jumped', 'over', 'a', 'lazy', 'dog']``

- No loops allowed! (We haven't talked about them yet.)
- Obviously the following approach gets zero points:

```ipython
# Simply assigning the final answer without list manipulation -- zero points!
list = ['the', 'quick', 'brown', 'fox', 'jumped', 'over', 'a', 'lazy', 'dog']
print(list)
```

<!-- grade_id: cell-34110f1a5e8dfb71, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d7728d4672965e61, points: 2, grade: true, solution: false -->
```python
assert new_list == ['the', 'quick', 'brown', 'fox', 'jumped', 'over', 'a', 'lazy', 'dog']
```

# Sandpit begins here
### Add your own cells below

```python

```