# Assigment 3

### Instructions
- Do NOT change the filename.
- Do NOT add or copy cells before the "Sandpit". Only add your own cells after the "Sandpit".
- Enter your name below.

### Enter your name here
Name: CHANGE ME

# Task 3 (4 points)

Write a function ```is_sorted``` that takes a list as an argument and checks whether the list is sorted in ascending order or not. The print statements below should all return the expected output. Also write a docstring in the functions head - between the triple quotes - describing the input and the output of the function.

- Please note: You MUST write your own loop to do the comparison. Any "ready made" / built-in functions such as ``sorted()`` or ``list.sort()`` is NOT allowed!

```
    print(is_sorted([-3, 0.5, 100])) --> True
    print(is_sorted([2, 3, 5, 4.5])) --> False
    print(is_sorted(['a', 'c', 'e'])) --> True
    print(is_sorted(['a', 'e', 'c'])) --> False
```

It's OK to use loop(s) for this task.

<!-- grade_id: cell-9552dd6d8eb8b442, points: 0, grade: false, solution: true -->
```python
def is_sorted(list_to_sort):
    # YOUR CODE HERE
    raise NotImplementedError()
```

# Tests

The cell below contains unit tests for your solution.

Run the below cell (after running all the above cells with your code).

<!-- grade_id: cell-73c1b287466cff38, points: 1, grade: true, solution: false -->
```python
assert is_sorted([-3, 0.5, 100]) == True
```

<!-- grade_id: cell-ac98d546c94e697e, points: 1, grade: true, solution: false -->
```python
assert is_sorted(['a', 'e', 'c']) == False
```

<!-- grade_id: cell-c527e963011c1dd2, points: 1, grade: true, solution: false -->
```python
# Hidden tests
```

<!-- grade_id: cell-885add18fca84b88, points: 1, grade: true, solution: false -->
```python
# Hidden tests
```

# Sandpit begins here
### Add your own cells below